---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise. kjflksj lfjs ldfjsal flksj dflasj dfljs ldfkj lsjfkl asjfl sjlfj sl


`@instructions`
This is instruction part lskjdl sjfl jsldf jflks jdflksj dlfjls ajflasjfdljaskf

`@hint`
Use this hint lsdj flsjdf ljsl jsl jflajlj aldj l j

`@pre_exercise_code`
```{python}
Try this first
```

`@sample_code`
```{python}

```

`@solution`
```{python}
This is the answer
```

`@sct`
```{python}

```
